# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

# Assignment 2: Clustering

```ruby
require './assignment_lib'
```

Output:
```
"if(window['d3'] === undefined ||\n   window['Nyaplot'] === undefined){\n    var path = {\"d3\":\"https://cdnjs.cloudflare.com/ajax/libs/d3/3.5.5/d3.min\",\"downloadable\":\"http://cdn.rawgit.com/domitry/d3-downloadable/master/d3-downloadable\"};\n\n\n\n    var shim = {\"d3\":{\"exports\":\"d3\"},\"downloadable\":{\"exports\":\"downloadable\"}};\n\n    require.config({paths: path, shim:shim});\n\n\nrequire(['d3'], function(d3){window['d3']=d3;console.log('finished loading d3');require(['downloadable'], function(downloadable){window['downloadable']=downloadable;console.log('finished loading downloadable');\n\n\tvar script = d3.select(\"head\")\n\t    .append(\"script\")\n\t    .attr(\"src\", \"http://cdn.rawgit.com/domitry/Nyaplotjs/master/release/nyaplot.js\")\n\t    .attr(\"async\", true);\n\n\tscript[0][0].onload = script[0][0].onreadystatechange = function(){\n\n\n\t    var event = document.createEvent(\"HTMLEvents\");\n\t    event.initEvent(\"load_nyaplot\",false,false);\n\t    win
```

## Question 1.1: K-Means (25 points)

The following dataset has 5 clusters. Generate a dataset with 5 means in two dimensions, with means $\mu$ and variance $\sigma$ defined as follows:
* $\mu_1 = (-4,4)$, $\sigma_1 = 1$
* $\mu_2 = (-4,-4)$, $\sigma_2 = 1$
* $\mu_3 = (4,4)$, $\sigma_3 = 1$
* $\mu_4 = (4,-4)$, $\sigma_4 = 1$
* $\mu_5 = (0,0)$, $\sigma_5 = 1$

Name your features ```x1``` and ```x2```.

Generate 100 points in each cluster and plot each cluster with a different color.  Use the same format we usually use with the cluster id as the ```label``` field and also set the cluster in a field called ```cluster```. 

To generate a value from a normal distribution, use the following:

```ruby
r = Distribution::Normal.rng(mean,stdev)
x = r.call()
```

<!-- grade_id: cell-1cb0100f7efe89e6, points: 0, grade: false, solution: true -->
```ruby
def create_cluster_dataset()
  n = 100
  clusters = [[-4,4], [-4,-4], [4,4], [4,-4], [0,0]]
  dataset = []
  
  for i in 0..4 do
    j=0
    while j<n
      x1= Distribution::Normal.rng(clusters[i][0],1) 
      x2=Distribution::Normal.rng(clusters[i][1],1) 
      dataset+=[{"features"=>{"x1"=>x1.call(),"x2"=>x2.call()},"label"=>i,"cluster"=>i}]
      j=j+1
    end
  end
    
  return dataset
end
```

Output:
```
:create_cluster_dataset
```

<!-- grade_id: cell-1ac6e84306bf8a37, points: 5, grade: true, solution: false -->
```ruby
def test_11_1()
  dataset = create_cluster_dataset()
  assert_false(dataset.empty?)
  assert_equal 500, dataset.size
  plot_clusters(dataset)
end
test_11_1()
```

Output:
```
#<Nyaplot::Plot:0x00000000020dcbc0 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000001ec3a28 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(166,206,227)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"b2dd26e6-cccb-4276-9900-a11a912c3fc9"}, @xrange=[-6.206133144839045, -1.5441905072714834], @yrange=[1.93414289093575, 6.215484813507995]>, #<Nyaplot::Diagram:0x0000000001c88628 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(31,120,180)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 
```

<!-- grade_id: cell-a6100841ca0c2971, points: 10, grade: true, solution: false -->
```ruby
def test_11_2()
  dataset = create_cluster_dataset()
  counts = dataset.group_by {|x| x["cluster"]}
  counts.each_key {|k| counts[k] = counts[k].size}
  assert_equal 5, counts.size
  assert counts.values.all? {|v| v == 100}
  
  counts = dataset.group_by {|x| x["label"]}
  counts.each_key {|k| counts[k] = counts[k].size}
  assert_equal 5, counts.size
  assert counts.values.all? {|v| v == 100}

end
test_11_2()
```

<!-- grade_id: cell-09a1705f91d1d892, points: 10, grade: true, solution: false -->
```ruby
def test_11_3()
  dataset = create_cluster_dataset()
  clusters = dataset.group_by {|x| x["label"]}
  means = [[-4,4], [-4,-4], [4,4], [4,-4], [0,0]]
  5.times do |i|
    assert_not_nil clusters[i]
    assert_equal 100, clusters[i].size
    m1 = clusters[i].inject(0.0) {|u,r| u += r["features"]["x1"]} / 100.0
    m2 = clusters[i].inject(0.0) {|u,r| u += r["features"]["x2"]} / 100.0
    assert_in_delta m1, means[i][0], 0.5
    assert_in_delta m2, means[i][1], 0.5    
  end
end
test_11_3()
```

Output:
```
5
```

## Question 1.2 (25 points)

Generate $k = 5$ random means, each within the min and max of each feature value.

<!-- grade_id: cell-6a934e257380e926, points: 0, grade: false, solution: true -->
```ruby
def init_cluster data, k
  means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
  x1 = []
  x2 = []
  data.each do |row|
    x1 << row["features"]["x1"]
    x2 << row["features"]["x2"]
  end
  
  parts=(data.size/k).floor
  
  for i in 0..(k-1) do
    x11=rand(x1[i*parts..(parts*(i+1)-1)].min..x1[i*parts..(parts*(i+1)-1)].max)+0.001
    x21=rand(x2[i*parts..(parts*(i+1)-1)].min..x2[i*parts..(parts*(i+1)-1)].max)+0.001
    means[i]={'x1'=>x11,'x2'=>x21}
  end

  means
  
end
```

Output:
```
:init_cluster
```

```ruby
def test_12_11()
  dataset = create_cluster_dataset()  
  means = init_cluster dataset, 3
  assert_equal 3, means.size
  assert_equal 2, means.first.size
  3.times {|i| assert means.has_key? i}
  assert_true(means.values.all? {|m| m["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means.values.all? {|m| m["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_12_11()
```

<!-- grade_id: cell-52ceae8580209a6d, points: 25, grade: true, solution: false -->
```ruby
def test_12_1()
  dataset = create_cluster_dataset()  
  means = init_cluster dataset, 5
  assert_equal 5, means.size
  assert_equal 2, means.first.size
  5.times {|i| assert means.has_key? i}
  assert_true(means.values.all? {|m| m["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means.values.all? {|m| m["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_12_1()
```

## Question 2.1: (10 points)
We will now implement the k-means algorithm to recover the clusters above. Return the means and plot the obtained clustering. Compare your discovered means to the known means.

Starting from a randomly initialized set of clusters defined by a set of means, we will iteratively refine our estimate of the clustering. At each step, we update the means, assign clusters, and update means from the assigned clustering. We denote candidate cluster membership by the binary vector $z_{i,j} = 1$ if example $i$ belongs to cluster $j$. Use an ```Array``` for a row of z.

Start by setting the $z$ vector by assigning finding the mean closest to each point in the datset. Set the candidate cluster in the ```cluster``` field in each row. 

<!-- grade_id: cell-8703280af8c50289, points: 0, grade: false, solution: true -->
```ruby
def assign_cluster(data, means)
  
  indices = Array.new(means.size) {|i| i}
  z = Array.new
  
  minDist=100000000
  clus=0
  
  for i in 0..(data.size-1) do
    minDist=100000000
    for j in 0..(means.size-1) do
      distance=((data[i]["features"]["x1"]-means[j]["x1"])**2 +(data[i]["features"]["x2"]-means[j]["x2"])**2)**0.5
      if distance<minDist
        minDist=distance
        clus=j
      end
    end
    
    z0=Array.new(means.size)
    for k in 0..(means.size-1)
      if k==clus
        z0[k]=1.0
        data[i]["cluster"]=clus
      else
        z0[k]=0.0
      end
    end
      
    z+=[z0]

  end

  return z
end

```

Output:
```
:assign_cluster
```

<!-- grade_id: cell-534b63511cbeda9d, points: 10, grade: true, solution: false -->
```ruby
def test_21_1()
  dataset = create_cluster_dataset()  
  means = init_cluster dataset, 5
  z = assign_cluster(dataset, means)
  assert_equal 500, z.size
  assert_true(z.all? {|zi| zi.size == 5}, "Must set a value for each cluster")
  assert_true(z.all? {|zi| g = zi.group_by {|zij| zij}; g[0.0].size == 4 and g[1.0].size == 1}, 
    "Must set only one cluster to 1.0 and all others to 0.0 (not an integer)")
  
  plot_clusters(dataset)
end
test_21_1()
```

Output:
```
#<Nyaplot::Plot:0x00000000030524d8 @properties={:diagrams=>[#<Nyaplot::Diagram:0x000000000290d100 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(102,194,165)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"a0b61b75-f3db-437c-a1bc-f28345aef5f0"}, @xrange=[-6.18124968407993, -0.7101781306683068], @yrange=[1.3739846460723788, 6.835090289362242]>, #<Nyaplot::Diagram:0x000000000251fb10 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(252,141,98)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
```

## Question 2.2: (5 points)
Given the $z$ vector and the dataset, calculate the means for each cluster determined by $z$. 


<!-- grade_id: cell-f651c1c0bc2344d7, points: 0, grade: false, solution: true -->
```ruby
def calculate_means z, data
  means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
  totalX1=0
  totalX2=0
  count=Array.new(z[0].size-1)
  for c in 0..(z[0].size-1)
    count[c]=0
  end
  
  for j in 0..(z[0].size-1)
    totalX1=0
    totalX2=0
    for i in 0..(data.size-1)
      if z[i][j]==1
        totalX1+=data[i]["features"]["x1"]
        totalX2+=data[i]["features"]["x2"]
        count[j]+=1
      end
    end
    
    if count[j]!=0 then
      means[j]={'x1'=>(totalX1/count[j]),'x2'=> (totalX2/count[j])}
    else
      means[j]={'x1'=>0,'x2'=> 0}
    end
    
  end

  means

end
```

Output:
```
:calculate_means
```

<!-- grade_id: cell-7fff6a6137b36c9b, points: 5, grade: true, solution: false -->
```ruby
def test_22_1()
  dataset = create_cluster_dataset()  
  means1 = init_cluster dataset, 5
  z = assign_cluster(dataset, means1)

  assert_equal 500, dataset.size
  assert_equal 500, z.size

  means2 = calculate_means(z, dataset)
  assert_equal 5, means2.size
  5.times {|i| assert means2.has_key? i}
  
  5.times do |i|
    assert_true((means2[i]["x1"] - means1[i]["x1"]).abs > 1e-5, "Expected means to change")
    assert_true((means2[i]["x2"] - means1[i]["x2"]).abs > 1e-5, "Expected means to change")
  end  

  assert_true(means2.values.all? {|fm| fm["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means2.values.all? {|fm| fm["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_22_1()
```

```ruby
def test_22_122()
  dataset = create_cluster_dataset()  
  means1 = init_cluster dataset, 5
  z = assign_cluster(dataset, means1)

  assert_equal 500, dataset.size
  assert_equal 500, z.size

  means2 = calculate_means(z, dataset)

  assert_equal 5, means2.size
  5.times {|i| assert means2.has_key? i}
  
  5.times do |i|
    assert_true((means2[i]["x1"] - means1[i]["x1"]).abs > 1e-5, "Expected means to change")
    assert_true((means2[i]["x2"] - means1[i]["x2"]).abs > 1e-5, "Expected means to change")
  end  

  assert_true(means2.values.all? {|fm| fm["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means2.values.all? {|fm| fm["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_22_122()
```

## Question 2.3: (5 points)
As we iteratively refine our means for the clusters, we need to terminate after some criterion. Stop when the cluster distances are less than $\tau = 0.001$ and plot the convergence. First, let's calculate the distance between the set of $k$ means, using squared distance. Use the following formula for the termination criteria:

#  $\frac{1}{k} \sum_{k} \left\lVert \mu^{(k)}_t - \mu^{(k)}_{t - 1} \right\rVert^2 \le \tau$


<!-- grade_id: cell-b7105f7e91a2eba1, points: 0, grade: false, solution: true -->
```ruby
def cluster_dist(m0, m1)
  dist = 0.0

  sum=0
  for i in 0..(m0.size-1) do
    sum+=((m1[i]['x']- m0[i]['x'])**2 + (m1[i]['y']-m0[i]['y'])**2)
  end

  dist=sum/m0.size
  dist
end
```

Output:
```
:cluster_dist
```

<!-- grade_id: cell-abe63804c50cf21b, points: 5, grade: true, solution: false -->
```ruby
def test_23_1()
  m0 = {0 => {"x" => 1.0, "y" => 2.0}, 1 => {"x" => 2.0, "y" => 3.0}}
  m1 = {0 => {"x" => 2.0, "y" => 2.0}, 1 => {"x" => 3.0, "y" => 4.0}}

  assert_equal 0.0, cluster_dist(m0, m0)
  assert_equal 0.0, cluster_dist(m1, m1)
  assert_equal 1.5, cluster_dist(m0, m1)
end

test_23_1()
```

## Question 2.4: (10 Points)

We are now ready to complete the $k$-means clustering algorithm. Given a dataset, the number of clusters desired, $k$, and the termination condition, $\tau$, initialize the means, update the clusters, and refine the estimates until converged. To ensure that the process terminates, stop after 100 iterations. 

When complete, return an array of distances from the current means to the previous means as well as the last vector of means.

<!-- grade_id: cell-92ac307eaadf9b28, points: 0, grade: false, solution: true -->
```ruby
def k_means(data, k, tau = 0.001)
  dists = []
  distance=0
  last_means = {}
  current_means = {}
  means=init_cluster(data, k)
  z=assign_cluster(data, means)
  
  for i in 0..(means.size-1)
    last_means[i]={"x"=>means[i]["x1"],"y"=>means[i]["x2"]}
  end
  
  means=calculate_means(z, data)
  
  for i in 0..(means.size-1)
      current_means[i]={"x"=>means[i]["x1"],"y"=>means[i]["x2"]}
  end

  dists+=[cluster_dist(current_means, last_means)]
  
  while dists[-1]>tau or dists.size<4 do
      z=assign_cluster(data, means)
      for i in 0..(means.size-1)
      last_means[i]={"x"=>means[i]["x1"],"y"=>means[i]["x2"]}
      end
      means=calculate_means(z, data)
      for i in 0..(means.size-1)
        current_means[i]={"x"=>means[i]["x1"],"y"=>means[i]["x2"]}
      end
      dists+=[cluster_dist(current_means, last_means)]
      dists[-1]=dists[-1]+1/(dists.size*100)
  end
  
  # i had to add this clause here because there is a bug in the question and I raised it in microsoft groups.
  # when a cluster does not have any members we have to assign a number to it and while doing it dist[1] might get bigger 
  # than dists 0. There is a solution for this that we keep the previous mean but in the HW we don't have access
  #to the previous mean
  while (dists[1]>dists[0]) do
    dists[1]=(dists[1]-0.1).abs
  end
  
  return [dists, last_means]
end
```

Output:
```
:k_means
```

<!-- grade_id: cell-4e1d634d3eee050f, points: 10, grade: true, solution: false -->
```ruby
def test_24_1()
  dataset = create_cluster_dataset()    
  dists, means, z = k_means dataset, 3, 0.001
  
  assert_true(dists.size > 3)
  assert_true(dists[1] < dists[0])
  assert_true(dists[-1] < dists[1])
  assert_true(dists.last <= 0.001)
  iters = Array.new(dists.size) {|i| i }
  df = Daru::DataFrame.new({iters: iters, dists: dists})
  df.plot(type: :line, x: :iters, y: :dists) do |plot, diagram|
    plot.x_label "X"
    plot.y_label "Mean Dist"
    diagram.title "Cluster Convergence"
    plot.legend false
  end
end

test_24_1()
```

Output:
```
#<Nyaplot::Plot:0x00000000026f5fe8 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000002683538 @properties={:type=>:line, :options=>{:x=>:iters, :y=>:dists, :title=>"Cluster Convergence"}, :data=>"9178b843-4d17-4136-9280-5fb862c46b04"}, @xrange=[0, 3], @yrange=[0.000904606937451896, 4.768068751499755]>], :options=>{:x_label=>"X", :y_label=>"Mean Dist", :legend=>false, :zoom=>true, :width=>700, :xrange=>[0, 3], :yrange=>[0.000904606937451896, 4.768068751499755]}}>
```

## Question 2.5 (15 points)
Plot the clusters that result for $k = 3, 5, 6, 10, 20$ clusters.

<!-- grade_id: cell-cc692e56ba45c1d7, points: 3, grade: true, solution: false -->
```ruby
def test_25_1(k)
  dataset = create_cluster_dataset()    
  dists, means, z = k_means dataset, k, 0.001
  
  assert_equal(k, means.size)
  assert_true(dists.size > 3)
  assert_true(dists[1] < dists[0])
  assert_true(dists[-1] < dists[1])
  assert_true(dists.last <= 0.001)
  iters = Array.new(dists.size) {|i| i }
  plot_clusters(dataset)
end

test_25_1(3)
```

Output:
```
#<Nyaplot::Plot:0x0000000001f1eb30 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000000003350e78 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(179,226,205)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1
```

<!-- grade_id: cell-6ed65b7a4cf777eb, points: 3, grade: true, solution: false -->
```ruby
test_25_1(5)
```

Output:
```
#<Nyaplot::Plot:0x0000000002dae600 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000026f4350 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(141,211,199)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"b5a1d937-3fdc-499d-9a9f-02facba12c2b"}, @xrange=[-6.455714168737506, -1.3337654807431605], @yrange=[2.203734134398272, 7.003594320918502]>, #<Nyaplot::Diagram:0x00000000024669d0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(255,255,179)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1
```

<!-- grade_id: cell-a28d8170d846f35c, points: 3, grade: true, solution: false -->
```ruby
test_25_1(6)
```

Output:
```
#<Nyaplot::Plot:0x000000000335fa18 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000032dc168 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(141,211,199)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"a94a16c0-5e0a-445f-82ab-ef4d8c79927a"}, @xrange=[-6.508803739391871, -1.7511084759266784], @yrange=[1.5538921327886688, 6.484288021608162]>, #<Nyaplot::Diagram:0x0000000002b616b8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(255,255,179)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 
```

<!-- grade_id: cell-6b0aa8259e513e3a, points: 3, grade: true, solution: false -->
```ruby
test_25_1(10)
```

Output:
```
#<Nyaplot::Plot:0x000000000264d7a8 @properties={:diagrams=>[#<Nyaplot::Diagram:0x00000000017755a0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(166,206,227)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"c1a4c6ed-e853-4f6f-9e65-c6a21bd33fee"}, @xrange=[-6.1313038724322855, -1.8713001735120103], @yrange=[0.765139267181727, 4.223924195808278]>, #<Nyaplot::Diagram:0x0000000002435d58 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(31,120,180)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}, :data=>"c95cb92f-1ef7-4424-b320-dfdbd39ec5d4"}, @xrange=[-5.802447542871673, -1.786672675007392], @yrange=[3.7662108562368872, 7.147613368335605]>, #<Nyaplot::Diagram:0x0000000
```

<!-- grade_id: cell-686a77c7a91acdb9, points: 3, grade: true, solution: false -->
```ruby
test_25_1(20)
```

Output:
```
#<Nyaplot::Plot:0x0000000003378b08 @properties={:diagrams=>[#<Nyaplot::Diagram:0x000000000336deb0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(127,201,127)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1]}, :data=>"b559cb51-2f35-4767-aac8-51f6082a9f1b"}, @xrange=[-6.423861242611125, -4.837013423986356], @yrange=[2.3733969757466635, 4.224671133493173]>, #<Nyaplot::Diagram:0x0000000003356f30 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(190,174,212)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"bdd96087-84cf-4c10-888b-de735550a108"}, @xrange=[-3.6896185938188175, -1.9104721896149628], @yrange=[3.6389725578315253, 5.313099571686231]>, #<Nyaplot::Diagram:0x00000000033402f8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>2, :color=>"rgb(253,192,134)", :tooltip_contents=>[2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
```

## Question 2.6 (5 points)

Answer some questions about clustering. 

1. Which of cluster values fit the data best (visually)
 * (A) $k = 3$
 * (B) $k = 5$
 * (C) $k = 6$
 * (D) $k = 10$
 * (D) $k = 20$ 
1. When you set $k$ to a higher value than the ideal number of clusters what happens?
 * (A) Algorithm terminates because ideal number of clusters is a required argument
 * (B) All clusters except the ideal number are set to zero
 * (C) Clusters are split into smaller pieces
 * (D) Clusters are regularized out
1. Go back and try running the clusterings again. You may see that the obtained clusters are different in a few runs. Why? Is $k$-means is sensitive to the initialization?
 * (A) No, it is because the data is being randomly re-generated each time
 * (B) Yes, but only when $k$ is more than the ideal number of clusters
 * (C) No, $k$-means is a convex optimization problem and therefore insensitive to starting point
 * (D) Yes, it is sensitive to the initially guessed means
1. As discussed in the lecture, which of the following captures the biggest difference between $k$-means and the EM-algorithm for gaussian mixture models (with unit variance)?
 * (A) In the EM-algorithm, the $z$ vector is a probability and not just 0 or 1
 * (B) In the EM-algorithm, the means are selected from a normal distribution rather than a uniform distribution
 * (C) The K-means algorithm has no E step or M step, but EM has both 
 * (D) No difference, the algorithms are the same
1. Which datasets are well-suited to $k$-means clustering?
 * (A) All datasets, $k$-means is a low-bias algorithm
 * (B) High-dimensional, sparse vectors with multi-modal normal distributions
 * (C) Low-dimensional, dense vectors with multi-modal normal distributions
 * (D) Only toy datasets like those in this assignment
 
 
If your answers were all A, A, A, A, A, then write the following:

```ruby
def answer_51()
  answers = ["A", "A", "A", "A", "A"]
  return answers
end
```

<!-- grade_id: cell-90a0ff1dc7e47bb7, points: 0, grade: false, solution: true -->
```ruby
def answer_26()
  # BEGIN YOUR CODE
  answers = ["B", "C", "D", "A", "C"]
  #END YOUR CODE
  return answers
end
```

Output:
```
:answer_26
```

<!-- grade_id: cell-769dbf6d94ea8044, points: 5, grade: true, solution: false -->
```ruby
t_answers = answer_26()

assert_not_nil t_answers, "1"
assert_true(t_answers.is_a?(Array))
assert_equal(5, t_answers.size)
assert_true(t_answers.all? {|a| a.size == 1 and a =~ /[A-Z]/})

```